```python

```

Output:
```
[11, 12, 22, 25, 34, 64, 90]
```

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Hai Dang Hoang"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
Line 7 decreases i by one unit to consider the preceding number. During the while loop, after we have identified that the number before the key is larger than the key (line 5), we push it up one position (line 6), then we consider the number before that (line 7) with the while loop again. We continue to do this until the number is no longer larger than the key, when we would insert the key (line 8). Therefore, line 7 is a necessary step to insert the key into the right position.

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i = j - 1
        while i>=0 and A[i]>key:
            A[i+1] = A[i]
            i = i-1
        A[i+1] = key #this is the error - the insertion must be outside of the while loop, or the final result would just be a series of the same number (all the "keys")
    return(A) #we must have the result returned to us as well
    raise NotImplementedError()

```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
We can try with different examples of A (I tried one and it works). The test code in the above cell also did not return an error.

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    n = len(A)
    for i in range(n):
        for j in range(0,n-i-1):
            if A[j]>A[j+1]:
                A[j],A[j+1]=A[j+1],A[j]
    return(A)
    raise NotImplementedError()
A  = [1,4,3,6,2,4,4,8,346,81]
bubbleSort(A)
```

Output:
```
[1, 2, 3, 4, 4, 4, 6, 8, 81, 346]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

```python
A = [5,6,3,6,34,665,7,2,1]
bubbleSort(A)
```

Output:
```
[1, 2, 3, 5, 6, 6, 7, 34, 665]
```

```python
A = [6,23,4,7,3,2,7,2,2,21,3,5]
bubbleSort(A)
```

Output:
```
[2, 2, 2, 3, 3, 4, 5, 6, 7, 7, 21, 23]
```

```python

```

```python

```

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
It is not an unambiguous proof - it is possible that for other examples of A, the code could be wrong. Before I arrived at the current version of the code, I had a version that worked with certain examples of A, but not others. Though this new version solved that problem, we cannot be completely sure that it is absolutely correct.

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    for i in range(0,len(A)-1):
        minidx = i
        for j in range (i+1, len(A)):
            if A[j] < A[minidx]:
                A[j], A[minidx] = A[minidx], A[j]
    return A
    raise NotImplementedError()
A = [6,23,4,7,3,2,7,2,2,21,3,5]
selectionSort(A)
```

Output:
```
[2, 2, 2, 3, 3, 4, 5, 6, 7, 7, 21, 23]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**